### Assignment Description

In this programming assignment, we will explore an agent in gridworld that uses value update to estimate the best action in a state.

First, we import the required libraries.

#### Note: Please do not modify any pre-defined variables. Doing so can affect the autograder results.

```python
##importing modules 
!pip install --user gym==0.20.0
import numpy as np
import pprint
from gridworld import GridworldEnv

pp = pprint.PrettyPrinter(indent=2)
```

Now we instantiate a gridworld environment from Sutton's Reinforcement Learning book chapter 4. 

You are an agent on an MxN grid, and your goal is to reach the terminal
state at the top left or the bottom right corner.

For example, a 4x4 grid looks as follows:

        T  o  o  o
        o  x  o  o
        o  o  o  o
        o  o  o  T

x is your position, and Ts are the two terminal states.

You can take actions in each direction (UP=0, RIGHT=1, DOWN=2, LEFT=3).

You do not need to implement anything for it, but it will help you assess and take action through the grid. Note that

- `env.P` represents the transition probabilities of the environment.
- `env.P[s][a]` is a list of transition tuples (prob, next_state, reward, done).
- `env.nS` is a number of states in the environment, 4 x 4 by default
- `env.nA` is a number of actions in the environment, i.e. (UP=0, RIGHT=1, DOWN=2, LEFT=3).

```python
env = GridworldEnv()  ###initializing the environments
print("Environment Loaded")
```

Now we create a helper function `action_estimation` to calculate the action value for all actions in a given state.

The arguments for the function are:
- `env`: Grid world environment
- `state`: The state to consider (int)
- `V`: The value to use as an estimator, vector of length `env.nS`
- `discount_factor`: Future discount factor

The function should return:
- A vector of length `env.nA` containing the expected value of each action.

> HINT: How can the Bellman Update equation be used here?

<!-- grade_id: cell-2f158181b532578b, points: 0, grade: false, solution: true -->
```python
##action estimation function

def action_estimation(env, state, V, discount_factor):
    A = np.zeros(env.nA)
    for a in range(env.nA):
        for prob, next_state, reward, done in env.P[state][a]:
            
            # START CODING HERE
        
            # your code here
            

            # END CODING HERE

    return A
```

Now we create a function `get_best_action` to find the best action and calculate the best action value.

The arguments for the function are:
- A vector of length `env.nA` containing the expected value of each action.

The function should return a tuple with two values:
- Best action `a`
- Action value for `a`

<!-- grade_id: cell-79ef4f20fc164650, points: 0, grade: false, solution: true -->
```python
def get_best_action(A):
    
    #NOTE: If there are multiple states with equal maximum value, choose the first one in order UP, RIGHT, DOWN, LEFT
    # For exampe if the values are [0, 2, 2, 1], your function should pick the direction as RIGHT as it appears before DOWN.
    
    # START CODING HERE
        
    # your code here
    

    # END CODING HERE    
```

Now create a function to create a deterministic policy using the optimal value function.

The arguments for the function are:
- `env`: Grid world environment
- `V`: The value to use as an estimator, vector of length `env.nS`
- `discount_factor`: Future discount factor

The function should return a matrix `policy` of size `env.nS` x `env.nA`

> Hint: A deterministic `policy` always selects one action per state for all `env.nS` states.

> HINT: What should be the probability of best_action in the policy?

<!-- grade_id: cell-2dc9d98af1ca17cf, points: 0, grade: false, solution: true -->
```python
def policy_estimation(env, V, discount_factor):    
    policy = np.zeros([env.nS, env.nA])
    
    for s in range(env.nS):
        # find the best action for this state
        A = action_estimation(env, s, V, discount_factor)
        
        best_action, best_action_value = get_best_action(A)

        # START CODING HERE
        
        # your code here
        

        # END CODING HERE  
        
    return policy
```

Finally, we implement the value iteration algorithm.

The arguments for `value_iteration` are:
- `env`: Gridworld environment
- `theta`: We stop evaluation once our value function change is less than theta for all states
- `discount_factor`: (Gamma) Future discount factor
- `max_iterations`: A constant limit on how many times the update should be done

And it should return:
- A tuple (`policy`, `V`) of the optimal policy and the optimal value function.

```python
def value_iteration(env, theta=0.0001, discount_factor=1.0, max_iterations=1000):
    V = np.zeros(env.nS)
    
    for iteration in range(max_iterations):
        # Stopping condition
        delta = 0
        
        # Update each state...
        for s in range(env.nS):
            # Do a one-step lookahead to find the best action
            A = action_estimation(env, s, V, discount_factor)
            
            best_action, best_action_value = get_best_action(A)
            
            # Calculate delta across all states seen so far
            delta = max(delta, np.abs(best_action_value - V[s]))
            
            # Update the value function. Ref: Sutton book eq. 4.10. 
            V[s] = best_action_value        
            
        # Check if we can stop 
        if delta < theta:
            break
    
    policy = policy_estimation(env, V, discount_factor)
    
    return policy, V
```

Now let's check the results of `value_iteration`.

```python
policy, v = value_iteration(env)

print("Policy Probability Distribution:")
print(policy)
print("")

print("Reshaped Grid Policy (0=up, 1=right, 2=down, 3=left):")
print(np.reshape(np.argmax(policy, axis=1), env.shape))
print("")

print("Value Function:")
print(v)
print("")

print("Reshaped Grid Value Function:")
print(v.reshape(env.shape))
print("")
```

If you have implemented the aforementioned functions correctly, the following test should pass without issue.

<!-- grade_id: cell-bfdf444f701e05f6, points: 10, grade: true, solution: false -->
```python
# Test the value function
expected_v = np.array([ 0, -1, -2, -3, -1, -2, -3, -2, -2, -3, -2, -1, -3, -2, -1,  0])
assert np.allclose(v, expected_v, rtol=1e-4) 
```

You are all set.

We run some hidden test cases on your code. You should be able to see the results once you submit the assignment.

<!-- grade_id: cell-624a8278382972c0, points: 15, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-87d7c445045b467b, points: 15, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-6f31fa320754ae14, points: 15, grade: true, solution: false -->
```python

```

```python

```